# Exercise Set 1

The purpose of this notebook is to walk you through some basics of programming in Python. If you have had any experience with programming languagues in the past, this is very likely to be a review, but if you have not, this will provide fundamental programming concepts that will be essential to using Python in a data science context.

The goal of this set of exercises is not to make you an expert programmer. It is to give a brief "grammar" lesson in the language of Python. The experiential assignments we do in class will do much more than any exercise in building your "conversational" abilities in Python, but having a brief introduction to the grammer will make the "conversation time" in class more productive.

If anything in this introduction is not clear, I would encourage you to google the relevant concept. There is an abundance of excellent material on python freely available online.

As you are going through this notebook, I encourage you to add cells wherever you feel like to play around with some of the concepts. Adding cells will not mess up the problems at the end, and trying to play with the code is a great way to test yourself in order to make sure you are following along.

Traditionally, the very first thing a new programmer does in a new language is write a "hello, world" program. This involves getting the computer to print the phrase "Hello, World!". In Python, this is very easy. All "strings" (i.e. anything that is not made up exclusively of numbers) are quoted with either double `"` or single `'` quotes. To display a string, you use the `print` function. Therefore, the complete program can be written in the single line below:

```python
print("Hello, World!")
```

Output:
```
Hello, World!
```

Now we are ready for some more interesting material.

## Acknowledgements

Significant portions of this exercise has been compiled using material from a number of open source python tutorials. I'm deeply indebted to their authors for examples and explanations. The tutorials that were used in compiling this exercise are:

1. [Python Practice](http://python.berkeley.edu/learn/)
2. [Lectures on Scientific Python](https://github.com/jrjohansson/scientific-python-lectures)
3. [Python Lectures](https://github.com/rajathkmp/Python-Lectures)


As well as various blog posts and Stack Overflow questions.

## Variables or Names for Stuff

A name that is used to denote something or a value is called a variable. In python, variables can be declared and values can be assigned to it as follows,

```python
# This is a comment. Anything after a "#" sign in a code block is not run.
x = 2     # This is also a comment
y = 5
xy = 'Hey'
```

```python
print(x+y)
print(xy)
# We can also print multiple things at a time.
print(x+y, xy)
```

Output:
```
7
Hey
7 Hey
```

Multiple variables can be assigned with the same value.

```python
x = y = 1
```

```python
print(x,y)
```

Output:
```
1 1
```

Variables are stored in _memory_, which is not persistent. That means that every time the "kernel" of the python process shuts down, it forgets all variables. The kernel will shut down for a variety of reasons. For example, if you shut down you JupyterHub server, the kernel will shut down. If you run out of memory, the kernel shuts down. Sometimes you restart the kernel becuase something is funny with your notebook.

Anytime this happens, you will lose all variables, and you will have to recreate them by running the relevant code again. For the purposes of this class, this will mostly be fine. It won't take that long to recreate variables, so you shouldn't worry about having to do this. Generally, it is good practice to structure your notebook such that you can just hit "Run All Cells" when you start it up, and it gets you to wherever you need to be to keep on with the project.

We will look at ways to save things to disk, so that it is persisent across kernel shutdowns, but that will not be until later in the course when runtimes get long enough that saving such things is necessary.

## Operators

Operators are generally symbols that act on variables or pieces of data. The structure is typically `variable operator other_variable`, much like formulas in math. Below are common operators.

### Arithmetic Operators

| Symbol | Task Performed |
|----|---|
| +  | Addition |
| -  | Subtraction |
| /  | division |
| %  | mod |
| *  | multiplication |
| //  | floor division |
| **  | to the power of |

```python
1+2
```

Output:
```
3
```

```python
2-1
```

Output:
```
1
```

```python
1*2
```

Output:
```
2
```

```python
1/2
```

Output:
```
0.5
```

```python
15%10
```

Output:
```
5
```

Floor division is division that rounds down to the nearest integer.

```python
2.8//2.0
```

Output:
```
1.0
```

Standard math operators work as expected on numbers.

```python
a = 2
b = 3
print(a + b)
print(a * b)
print(a ** b)  # a to the power of b (a^b does something completely different!)
print(a / b)
```

Output:
```
5
6
8
0.6666666666666666
```

### String Operators

Strings also have operators. You can add them:

```python
print('Hello,' + ' ' + 'World!')
```

Output:
```
Hello, World!
```

You can even multiply them:

```python
print("hello"*3)
```

Output:
```
hellohellohello
```

### Relational Operators

Something that we will use a lot are relational operators. They compare things and return either `True` or `False`, and we can use these truth values to do various things.

| Symbol | Task Performed |
|----|---|
| == | True, if it is equal |
| !=  | True, if not equal to |
| < | less than |
| > | greater than |
| <=  | less than or equal to |
| >=  | greater than or equal to |

```python
z = 1
```

Note that `=` and `==` are different things. `=` sets a variable equal to something. `==` tests if two things are equal, but it doesn't set anything equal.

```python
z == 1
```

Output:
```
True
```

```python
z == 0
```

Output:
```
False
```

```python
print(z)
```

Output:
```
1
```

```python
z > 1
```

Output:
```
False
```

```python
z >= 1
```

Output:
```
True
```

Boolean operators take two boolean values (either `True` or `False`) and return a boolean values for two variables a and b, the resulting output for the `and` operator is,

| a | b | a and b |
|----|---|---|
| `True` | `True` | `True` |
| `False`  | `True` | `False` |
| `True` | `False` | `False` |
| `False` | `False` | `False` |

I.e., it's only `True` if both inputs are `True`. For the `or` operator, it is `True` if at least one of the inputs is `True`:

| a | b | a or b |
|----|---|---|
| `True` | `True` | `True` |
| `False`  | `True` | `True` |
| `True` | `False` | `True` |
| `False` | `False` | `False` |

A simple but important boolean operator is `not`,

| a | not a |
|----|---|
| `True` | `False` |
| `False`  | `True` |

```python
a = (1 > 3)
b = (3 == 3)
print("a is " + str(a))
print("b is " + str(b))
print("a or b is " + str(a or b))
print("a and b is " + str(a and b))
print("not a is " + str(not a))
```

Output:
```
a is False
b is True
a or b is True
a and b is False
not a is True
```

Why did I use `str` when I wanted to print things out? Remember, we can add strings together. Let's see what happens when we don't use it.

```python
print("a is " + a)
```

Output:
```
TypeError                                 Traceback (most recent call last)
<ipython-input-87-051ae377a1ea> in <module>
----> 1 print("a is " + a)

TypeError: can only concatenate str (not "bool") to str
```

This is our first example of a "traceback error". You will likely see these a lot while programming, and they mean that there is something wrong with the code. We will practice deciphering them, but this one says `TypeError: can only concatenate str (not "bool") to str`. That is becuase we are trying to "add" two things together that aren't the same. `a` is a "bool" (or a True or False value, a.k.a. Boolean), and `"a is "` is a str (or string). We can confirm that with the `type` function.

```python
print(type("a is "))
print(type(a))
```

While we can add two strings, together, we can't add a string and a bool. In order to do this, we have to first convert `a` into a string. We do this using the `str` function. So, what is a function?

## Functions

These will be very familiar to anyone who has programmed in any language or used excel extensively, and work like you
would expect.

```python
# There are thousands of functions that operate on things
print(type(3))
print(len('hello'))
print(round(3.3))
```

__TIP:__ To find out what a function does, you can type it's name and then a question mark to
get a pop up help window.

```python
?round
round(3.14159, 2)
```

__TIP:__ Many useful functions are not in the Python built in library, but are in external
scientific packages. These need to be imported into your Python notebook (or program) before
they can be used.

```python
import numpy as np
```

These functions can then be accessed by using `np.FUNCTION_NAME`. Some examples of numpy functions and "things":

```python
print(np.sqrt(4))
print(np.pi)  # Not a function, just a variable
print(np.sin(np.pi))
```

We can also import specific functions so that we don't have to use the prefix `np.` using the `from LIBRARY_NAME import FUNCTION_NAME` syntax.

```python
from numpy import pi
```

```python
print(pi)
```

## Methods

Before we get any farther into the Python language, we have to say a word about "objects". We
will not be teaching object oriented programming in this class, but you will encounter objects
throughout Python (in fact, even seemingly simple things like numbers and strings are actually
objects in Python).

In the simplest terms, you can think of an object as a small bundled "thing" that contains within
itself both data and functions that operate on that data. For example, strings in Python are
objects that contain a set of characters and also various functions that operate on the set of
characters. When bundled in an object, these functions are called "methods".

Instead of the "normal" `function(arguments)` syntax, methods are called using the
syntax `object.method(arguments)`.

```python
a = 'hello, world'
print(type(a))
```

The `class` above indicates it is an object.

Objects have bundled methods. For example:

```python
print(a.capitalize())
print(a.replace('l', 'X'))
```

You can combine operators and methods like so:

```python
((a + " y'all " + a)*3).replace('l', 'X').replace('dh', 'd h')
```

Let's break it down. We first create the string `"hello, world y'all hello, world"` using `a + " y'all " + a`. Then we multiply that string three times with `(a + " y'all " + a)*3` to give us `"hello, world y'all hello, worldhello, world y'all hello, worldhello, world y'all hello, world"`. Then we replace all `l`s with `X`s. Then we replace `dh` with `d h` to give us the spaces in between strings. While this may look a little complicated, it quickly becomes second nature to parse this kind of syntax.

## Collections or Putting Stuff in Other Stuff

We will regularly want to bundle together things. This will often be us bundling together data points. Though sometimes, we will bundle together steps in a machine learning model to package it all together.

### Lists

Lists are probably the handiest and most flexible type of collection. We will use them all of the time.

Lists are declared with square brackets `[]`. 

Individual elements of a list can be selected using the syntax `a[ind]`.

```python
# Lists are created with square bracket syntax
a = ['blueberry', 'strawberry', 'pineapple', 'orange']
print(a, type(a))
```

Lists (and all collections) are also indexed with square brackets. NOTE: The first index is zero, not one

```python
print(a[0])
print(a[1])
```

```python
## You can also count from the end of the list
print('last item is:', a[-1])
print('second to last item is:', a[-2])
```

```python
# you can access multiple items from a list by slicing, using a colon between indexes
# NOTE: The end value is not inclusive
print('a =', a)
print('get first two:', a[0:2])
print('get middle two:', a[1:3])
```

```python
# You can leave off the start or end if desired
print(a[:2])
print(a[2:])
print(a[:])
print(a[:-1])
```

```python
print (a)
# We can skip items by adding a third parameter for a skip frequency
print(a[1:4:2])
# We can reverse a by making the third parameter negative
print(a[::-1])
```

```python
# Lists are objects, like everything else, and have methods such as append
a.append('banana')
print(a)

a.append([1,2])
print(a)

a.pop()
print(a)
```

__TIP:__ A 'gotcha' for some new Python users is that many collections, including lists,
actually store pointers to data, not the data itself. 

Remember when we set `b=a` and then changed `a`?

What happens when we do this in a list?

```python
a = 1
b = a
a = 2
## What is b?
print('What is b?', b)

a = [1, 2, 3]
b = a
print('original b', b)
a[0] = 42
print('What is b after we change a ?', b)
```

In order to avoid this, you have to `copy` the list.

```python
a = [1, 2, 3]
b = a.copy() # We set b equal to a copy of a
print('original b', b)
a[0] = 42
print('What is b after we change a ?', b)
print('What is a ?', a)
```

Often times, we will need to know how long a list is. We can get that with `len`:

```python
print(len(a))
```

### Tuples

We won't say a whole lot about tuples except to mention that they basically work just like lists, with
two major exceptions:

1. You declare tuples using `()` instead of `[]`
1. Once you make a tuple, you can't change what's in it (referred to as immutable)

You'll see tuples come up throughout the Python language, and over time you'll develop a feel for when
to use them. Often times, they are used as arguments to functions.

In general, they're often used instead of lists:

1. to group items when the position in the collection is critical, such as coord = (x,y)
1. when you want to prevent accidental modification of the items, e.g. shape = (12,23)

```python
xy = (23, 45)
print(xy[0])
xy[0] = "this won't work with a tuple"
```

### Anatomy of a traceback error

Another error! Let's take a closer look at how to read it.

Traceback errors are `raised` when you try to do something with code it isn't meant to do.  It is also meant to be informative, but like many things, it is not always as informative as we would like.

Looking at our error:

``` python
    TypeError                                 Traceback (most recent call last)
    <ipython-input-25-4d15943dd557> in <module>()
          1 xy = (23, 45)
          2 xy[0]
    ----> 3 xy[0] = 'this wont work with a tuple'

    TypeError: 'tuple' object does not support item assignment
```
    
1. The command you tried to run raise a **TypeError**  This suggests you are using a variable in a way that its **Type** doesnt support
2. the arrow ----> points to the line where the error occurred, In this case on line 3 of your code form the above line.
3. Learning how to **read** a traceback error is an important skill to develop, and helps you know how to ask questions about what has gone wrong in your code. Once you can read the error, you will know a) what line is wrong with your code, and b) what to google to try to fix it.

They can look a little more complicated, as the below piece of code shows:

```python
from sklearn.impute import SimpleImputer
SimpleImputer(missing_values=np.nan, strategy='mean').fit([[0, 1, 1, "nonsense", np.nan, np.nan]])
```

Here, there are multiple `--->` arrows in the error. That's becuase this code uses a method, that calls other code, and the error actually happens several steps down. However, the error is due to the method's input `[[0, 1, 1, "nonsense", np.nan, np.nan]]`, and we can see this, becuase one of the arrows points to 

``` python
----> 2 SimpleImputer(missing_values=np.nan, strategy='mean').fit([[0, 1, 1, "nonsense", np.nan, np.nan]])
```

Since this is a line of our code, this is likely where the error is. Moreover, since the line at the bottom is:

``` python
AttributeError: 'list' object has no attribute 'dtype'
```

it probably has something to do with the list we are giving it as input. If we replace `"nonsense"` with a number, then the error goes away.

```python
SimpleImputer(missing_values=np.nan, strategy='mean').fit([[0, 1, 1, 1, np.nan, np.nan]])
```

Learning to read these errors takes a little time, but when you do get a sense for what it's saying, you will be much better equipped to understand why your code isn't working.

### Dictionaries

Dictionaries are the collection to use when you want to store and retrieve things by their names
(or some other kind of key) instead of by their position in the collection. A good example is a set
of model parameters, each of which has a name and a value. Dictionaries are declared using `{}`.

We are not going to use dictionaries a lot, but we will see them crop up. It is good to have a basic familiarity with them.

```python
# Make a dictionary of model parameters
convertors = {'inches_in_feet' : 12,
              'inches_in_metre' : 39}

print(convertors)
print(convertors['inches_in_feet'])
```

```python
## Add a new key:value pair
convertors['metres_in_mile'] = 1609.34
print(convertors)
```

```python
# Raise a KEY error
print(convertors['blueberry'])
```

## Loops or How to Repeat Yourself

So far, everything that we've done could, in principle, be done by hand calculation. In this section
and the next, we really start to take advantage of the power of programming languages to do things
for us automatically.

We start here with ways to repeat yourself. The two most common ways of doing this are known as `for`
loops and `while` loops. For loops in Python are useful when you want to cycle over all of the items
in a collection (such as all of the elements of an array or a list), and while loops are useful when you want to
cycle for an indefinite amount of time until some condition is met.

The basic examples below will work for looping over lists, tuples, and arrays. Looping over dictionaries
is a bit different, since there is a key and a value for each item in a dictionary. Have a look at the
Python docs for more information.

```python
# A basic for loop - don't forget the white space!
wordlist = ['hi', 'hello', 'bye']
for word in wordlist:
    print(word + '!')
```

**Note on indentation**: Notice the indentation once we enter the `for` loop.  Every idented statement after the `for` loop declaration is part of the `for` loop.  This rule holds true for `while` loops, `if` statements, functions, etc. Required identation is one of the reasons Python is such a beautiful language to read.

If you do not have consistent indentation you will get an `IndentationError`.  Fortunately, most code editors will ensure your indentation is correction.

__NOTE__ In Python the default is to use four (4) spaces for each indentation. This will happen automatically in a Jupyter Notebook if you hit tab at a place where you can indent things.

```python
# Indentation error: Fix it!
for word in wordlist:
    new_word = word.capitalize()
    print(new_word + '!') # Bad indent
```

```python
# Sum all of the values in a collection using a for loop
numlist = [1, 4, 77, 3]

for num in numlist:
    total = total + num
    
print("Sum is", total)
```

```python
# Often we want to loop over the indexes of a collection, not just the items
print(wordlist)

for i, word in enumerate(wordlist):
    print(i, word, wordlist[i])
```

`while` loops are useful when you don't know how many steps you will need, and want to stop once a certain condition is met. They keep going until a condition is `False` (remember those boolean operators from before?).

```python
step = 0
prod = 1
while prod < 100:
    step = step + 1
    prod = prod * 2
    print(step, prod)
    
print('Reached a product of', prod, 'at step number', step)
```

We will often want to loop over a list of integers. We can do this with `range`.

```python
for i in range(10):
    print(i)
```

```python
for i in range(2,10):
    print(i)
```

```python
for i in range(2,10,4):
    print(i)
```

```python
?range
```

`range` does not actually give us a `list`, it gives us a `generator` which we can pull values from. However, we can't use it like a list.

```python
a = range(2,10)
print(a)
```

A list would have printed the values. To make it a list, we can use the `list` function.

```python
print(list(a))
print(list(range(2,10)))
```

## `if` Statements or Making Choices

Often we want to check if a condition is True and take one action if it is, and another action if the
condition is False. We can achieve this in Python with an if statement.

__TIP:__ You can use any expression that returns a boolean value (True or False) in an if statement.

```python
# A simple if statement.
# Try changing x and re-running this block, and see what happens.
x = 3
if x > 0:
    print('x is positive')
elif x < 0:
    print('x is negative')
else:
    print('x is zero')
```

```python
# If statements can rely on boolean variables
x = -1
test = (x > 0)
print(type(test))
print(test)

if test:
    print('Test was true')
```

## Defining Functions or Learning How Not to Repeat Yourself

One way to write a program is to simply string together commands, like the ones described above, in a long
file, and then to run that file to generate your results. This may work, but it can be cognitively difficult
to follow the logic of programs written in this style. Also, it does not allow you to reuse your code
easily - for example, what if we wanted to run our logistic growth model for several different choices of
initial parameters?

The most important ways to "chunk" code into more manageable pieces is to create functions and then
to gather these functions into modules, and eventually packages. Below we will discuss how to create
functions and modules. A third common type of "chunk" in Python is classes, but we will not be covering
object-oriented programming in this class.

```python
x = 3.333333
print(round(x, 2))
print(np.sin(x))
```

It is very easy to write your own functions.

```python
# It's very easy to write your own functions
def multiply(x, y):
    z = x*y
    return z
```

Note that in the above, the function `return`s the value `z`. A function will often, `return` a value, and this allows you to do things like write `answer = multiply(5, 4)` and have the variable `answer` return `20`. If you do not have the `return` line, then you if you wrote `answer = multiply(5, 4)`, while it wouldn't give you an error, the variable `answer` would not have anything in it.

```python
# Once a function is "run" and saved in memory, it's available just like any other function
print(type(multiply))
print(multiply(4, 3))
```

```python
# It's useful to include docstrings to describe what your function does
def say_hello(time, people):
    '''
    Function says a greeting. Useful for engendering goodwill
    '''
    return 'Good ' + time + ', ' + people
```

**Docstrings**: A docstring is a special type of comment that tells you what a function does. It goes between `'''` and `'''` in the beginning of a function.  You can see them when you ask for help about a function. This docstring is completely optional, and we will often define functions that do not have it.

```python
say_hello('afternoon', 'friends')
```

```python
?say_hello
```

```python
# All arguments must be present, or the function will return an error
say_hello('afternoon')
```

```python
# Keyword arguments can be used to make some arguments optional by giving them a default value
# All mandatory arguments must come first, in order
def say_hello(time, people='friends'):
    return 'Good ' + time + ', ' + people
```

```python
say_hello('afternoon')
```

```python
say_hello('afternoon', 'students')
```

You can put `if` statements, `for` loops, `while` loops, or any other piece of python code in a function. Note that the indentation for `for` loops and other things is nested within the indentation for the function.

```python
def sum_list(numlist):
    total = 0
    for num in numlist:
        total = total + num
    return total
```

```python
sum_list([1, 2, 10, 10])
```

Note that in the above function `sum_list`, there is a variable called `total`. It is only accessbile within that function. If you try to print it outside of the function look at what happens:

```python
print(total)
```

Moreover, `total` is reset everytime that function is used, so it is kind of a way to have a temporary variable that is only used in that function and then thrown away. This can often be very useful to keep track of things in a function for temporary purposes.

## Exercises

Below are 10 problems, each worth 5 points each for a total of 50 points. Most of the problems (except for the first two) require you to complete a function that I start in order to get the function to output the specified thing given the input. Remember when you are writing functions about how code needs to be indented. Also, after each problem, there are tests associated with your answer. If you can run the cell after your answer and not get any errors, then you very likely have gotten the question right. If you have errors, hopefully the error will help you identify the mistake.

Note that just because you don't get errors doesn't mean that you got the question correct. I have some additional tests held back that I do not show here, though if you pass the ones shown, you will likely pass those as well.

**TIP:** For the problems that ask you to write your code in a function (problems 3-10), you will need to make sure that your function `return`s the answer. Do not just `print()` your answer. While it will look like your function is doing the right thing when you run it by itself, the function will not actually be returning any answer at all and the grading cell will give an error and count it as wrong.

When you are done with the assignment, just save the notebook with it's original name. It will be automatically collected at the due date and time.

#### Problem \#1 - 5 points

Create a variable `answer_1` and assign to it the value `10`.

<!-- grade_id: problem_1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
answer_1 = 10
print(answer_1)
```

Output:
```
10
```

<!-- grade_id: problem_1_test, points: 5, grade: true, solution: false -->
```python
# THIS IS A GRADING CELL. DO NOT EDIT.
print(answer_1)
```

Output:
```
10
```

#### Problem \#2 - 5 points

Store in the variable `answer_2` a list of integers from `15` to `245` inclusive (i.e., the first number in the list should be `15` and the last should be `245`). So, your code will look something like
``` python
answer_2 = ...
```
I would suggest that you confirm that the answer is what you expect through printing it out. Note that you are free to just type in this whole list, but I recommend you look for an easier way.

<!-- grade_id: problem_2, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
answer_2 = list(range(15,246))
print(answer_2)
```

Output:
```
[15, 16, 17, 18, 19, 20, 21, 22, 23, 24, 25, 26, 27, 28, 29, 30, 31, 32, 33, 34, 35, 36, 37, 38, 39, 40, 41, 42, 43, 44, 45, 46, 47, 48, 49, 50, 51, 52, 53, 54, 55, 56, 57, 58, 59, 60, 61, 62, 63, 64, 65, 66, 67, 68, 69, 70, 71, 72, 73, 74, 75, 76, 77, 78, 79, 80, 81, 82, 83, 84, 85, 86, 87, 88, 89, 90, 91, 92, 93, 94, 95, 96, 97, 98, 99, 100, 101, 102, 103, 104, 105, 106, 107, 108, 109, 110, 111, 112, 113, 114, 115, 116, 117, 118, 119, 120, 121, 122, 123, 124, 125, 126, 127, 128, 129, 130, 131, 132, 133, 134, 135, 136, 137, 138, 139, 140, 141, 142, 143, 144, 145, 146, 147, 148, 149, 150, 151, 152, 153, 154, 155, 156, 157, 158, 159, 160, 161, 162, 163, 164, 165, 166, 167, 168, 169, 170, 171, 172, 173, 174, 175, 176, 177, 178, 179, 180, 181, 182, 183, 184, 185, 186, 187, 188, 189, 190, 191, 192, 193, 194, 195, 196, 197, 198, 199, 200, 201, 202, 203, 204, 205, 206, 207, 208, 209, 210, 211, 212, 213, 214, 215, 216, 217, 218, 219, 220, 221, 222, 223, 224, 225, 226, 227, 228, 229, 230, 231,
```

<!-- grade_id: problem_2_test, points: 5, grade: true, solution: false -->
```python
# THIS IS A GRADING CELL. DO NOT EDIT.
print(answer_2)
```

Output:
```
[15, 16, 17, 18, 19, 20, 21, 22, 23, 24, 25, 26, 27, 28, 29, 30, 31, 32, 33, 34, 35, 36, 37, 38, 39, 40, 41, 42, 43, 44, 45, 46, 47, 48, 49, 50, 51, 52, 53, 54, 55, 56, 57, 58, 59, 60, 61, 62, 63, 64, 65, 66, 67, 68, 69, 70, 71, 72, 73, 74, 75, 76, 77, 78, 79, 80, 81, 82, 83, 84, 85, 86, 87, 88, 89, 90, 91, 92, 93, 94, 95, 96, 97, 98, 99, 100, 101, 102, 103, 104, 105, 106, 107, 108, 109, 110, 111, 112, 113, 114, 115, 116, 117, 118, 119, 120, 121, 122, 123, 124, 125, 126, 127, 128, 129, 130, 131, 132, 133, 134, 135, 136, 137, 138, 139, 140, 141, 142, 143, 144, 145, 146, 147, 148, 149, 150, 151, 152, 153, 154, 155, 156, 157, 158, 159, 160, 161, 162, 163, 164, 165, 166, 167, 168, 169, 170, 171, 172, 173, 174, 175, 176, 177, 178, 179, 180, 181, 182, 183, 184, 185, 186, 187, 188, 189, 190, 191, 192, 193, 194, 195, 196, 197, 198, 199, 200, 201, 202, 203, 204, 205, 206, 207, 208, 209, 210, 211, 212, 213, 214, 215, 216, 217, 218, 219, 220, 221, 222, 223, 224, 225, 226, 227, 228, 229, 230, 231,
```

#### Problem \#3 - 5 points

Write a function that given a string (called `string`, return the string repeated, but also replace all letters 'a' with 'b'. I.e., given the string "apple", it should return "bpplebpple".

<!-- grade_id: problem_3, points: 0, grade: false, solution: true -->
```python
def double_and_replace(string):
    # YOUR CODE HERE
    return(string*2).replace('a', 'b')

print(double_and_replace("apple"))
```

Output:
```
bpplebpple
```

<!-- grade_id: problem_3_test, points: 5, grade: true, solution: false -->
```python
# THIS IS A GRADING CELL. DO NOT EDIT.
from nose.tools import assert_equal
assert_equal(double_and_replace("apple"), "bpplebpple")
assert_equal(double_and_replace("test"), "testtest")
```

#### Problem \#4 - 5 points

Given two strings, a and b, write a function that returns the result of putting them together in the order abba, e.g. "Hi" and "Bye" returns "HiByeByeHi".

<!-- grade_id: problem_4, points: 0, grade: false, solution: true -->
```python
def make_abba(a, b):
    # YOUR CODE HERE
    return(a+b+b+a)

print(make_abba("Hi","Bye"))
```

Output:
```
HiByeByeHi
```

<!-- grade_id: problem_4_test, points: 5, grade: true, solution: false -->
```python
# THIS IS A GRADING CELL. DO NOT EDIT.
from nose.tools import assert_equal
assert_equal(make_abba("Hi", "Bye"), "HiByeByeHi")
assert_equal(make_abba("a", "b"), "abba")
```

#### Problem \#5 - 5 points

Given a list of numbers that is at least of length 2, write a function called `sum2` that returns the sum of the first 2 elements in the list. I.e.

``` python
sum2([1, 2, 3])  3
sum2([1, 1])  2
sum2([1, 1, 1, 1])  2
```

<!-- grade_id: problem_5, points: 0, grade: false, solution: true -->
```python
def sum2(num_list):
    # YOUR CODE HERE
    return(sum(num_list[0:2]))
    
print(sum2([1, 2, 3]))
```

Output:
```
3
```

<!-- grade_id: problem_5_test, points: 5, grade: true, solution: false -->
```python
# THIS IS A GRADING CELL. DO NOT EDIT.
from nose.tools import assert_equal
assert_equal(sum2([1, 2, 3]), 3)
assert_equal(sum2([1, 1]), 2)
assert_equal(sum2([1, 1, 1, 1]), 2)
```

#### Problem \#6 - 5 points

Write a function that takes an integer `n` and computes the squares of `1,...,n`. I.e., if `n=3`, the output should be the list `[1, 4, 9]`. If `n=5`, the output should be `[1, 4, 9, 16, 25]`.

<!-- grade_id: problem_6, points: 0, grade: false, solution: true -->
```python
def squares(n):
    """Compute squares"""
    # YOUR CODE HERE
    
    squares = [x**2 for x in range(1,n+1)]
    return(squares)

print(squares(4))
```

Output:
```
[1, 4, 9, 16]
```

<!-- grade_id: problem_6_tests, points: 5, grade: true, solution: false -->
```python
# THIS IS A GRADING CELL. DO NOT EDIT.
from nose.tools import assert_equal
assert_equal(squares(1), [1])
assert_equal(squares(3), [1, 4, 9])
assert_equal(squares(5), [1, 4, 9, 16, 25])
```

#### Problem #7 - 5 points

Write a function called `sorta_sum` that given 2 numbers, a and b, return their sum. However, sums in the range 10..19 inclusive, are forbidden, so in that case just return 20.

``` python
sorta_sum(3, 4)  7
sorta_sum(9, 4)  20
sorta_sum(10, 11)  21
```

<!-- grade_id: problem_7, points: 0, grade: false, solution: true -->
```python
def sorta_sum(a, b):
    # YOUR CODE HERE
    if a+b >9 and a+b <20:
        return(20)
    else:
        return(a+b)

print(sorta_sum(10,5))
```

Output:
```
20
```

<!-- grade_id: problem_7_test, points: 5, grade: true, solution: false -->
```python
# THIS IS A GRADING CELL. DO NOT EDIT.
from nose.tools import assert_equal
assert_equal(sorta_sum(3, 4), 7)
assert_equal(sorta_sum(9, 4), 20)
assert_equal(sorta_sum(10, 11), 21)
```

#### Problem \#8 - 5 points

Write a function that takes a list of numbers called `num_list`, and returns the list with only the numbers less than or equal to an optional parameter `upper_limit`. For example, if the list `[1, 7, 4, 12, 19, 5]` is the input list, and the optional parameter `upper_limit` is set to `10`, then it should return the list `[1, 7, 4, 5]`.

HINT: It will probably be helpful to use a new empty list and then add any element that meets the criteria to the list. Then return that new list. Remember the `.append` method from above.

<!-- grade_id: problem_8, points: 0, grade: false, solution: true -->
```python
def filter_list(num_list, upper_limit = 10):
    '''Filter a list of numbers by the upper_limit'''
    # YOUR CODE HERE
    A = [a for a in num_list if a < upper_limit]
    return A

print(filter_list([1, 7, 4, 12, 19, 5], upper_limit=10))
print(filter_list([1, 7, 4, 12, 19, 5]))
```

Output:
```
[1, 7, 4, 5]
[1, 7, 4, 5]
```

<!-- grade_id: problem_8_test, points: 5, grade: true, solution: false -->
```python
# THIS IS A GRADING CELL. DO NOT EDIT.
from nose.tools import assert_equal
assert_equal(filter_list([1, 7, 4, 12, 19, 5], upper_limit=10), [1, 7, 4, 5])
assert_equal(filter_list([1, 7, 4, 12, 19, 5]), [1, 7, 4, 5])
```

#### Problem \#9 - 5 points

Write a function called `flip_it_and_reverse_it` that takes a list called `input_list` and returns the reverse of the list. For example, the list `[1, 2, 3, 4]` would return `[4, 3, 2, 1]`. The list `["This", "is", "not", "a", "palindrome"]` would return `["palindrome", "a", "not", "is", "This"]`.

<!-- grade_id: problem_9, points: 0, grade: false, solution: true -->
```python
def flip_it_and_reverse_it(input_list):
    # YOUR CODE HERE
    return(input_list[::-1])
    
print(flip_it_and_reverse_it([1, 2, 3, 4]))
print(flip_it_and_reverse_it(["This", "is", "not", "a", "palindrome"]))
```

Output:
```
[4, 3, 2, 1]
['palindrome', 'a', 'not', 'is', 'This']
```

<!-- grade_id: problem_9_test, points: 5, grade: true, solution: false -->
```python
# THIS IS A GRADING CELL. DO NOT EDIT.
from nose.tools import assert_equal
assert_equal(flip_it_and_reverse_it([1, 2, 3, 4]), [4, 3, 2, 1])
assert_equal(flip_it_and_reverse_it(["This", "is", "not", "a", "palindrome"]), ["palindrome", "a", "not", "is", "This"])
```

#### Problem \#10 - 5 points

Below is a broken function called `fix_me`. You can define it just fine, but when you try to run it, you see that it won't run. Fix the function so that it returns `"this is not a broken function!"`.

<!-- grade_id: problem_10, points: 0, grade: false, solution: true -->
```python
def fix_me():
    x = ("this", " is", " a", " broken", " function!")
    # YOUR CODE HERE
    #x[2] = " not a"
    #return ''.join(x)
    
    fixed = list(x)
    fixed[2] = " not a"
    x = tuple(fixed)
    return ''.join(x)
    
print(fix_me())
```

Output:
```
this is not a broken function!
```

```python
fix_me()
```

Output:
```
'this is not a broken function!'
```

<!-- grade_id: problem_10_test, points: 5, grade: true, solution: false -->
```python
# THIS IS A GRADING CELL. DO NOT EDIT.
from nose.tools import assert_equal
assert_equal(fix_me(), "this is not a broken function!")
```